# Data Foundations -- Problem Set 3: Intro to Python¶
In this problem set, you will revisit some basic concepts that were first introduced in Module 2 in R. Now, you will be writing Python code to solve similar tasks. This problem set highlights how "coding is coding," meaning most times, you can solve the same problem with different coding languages. Often, it's your preference (or your team's) that determines which language you code in.

As always, use sources like Google, Stack Overflow, documentation, and other forums for help. Try to avoid AI in order to develop your own ability to conceptualize solution strategies.

## Question One (1 Point) 
Create two numeric variables named var_a and var_b and assign any numeric value to them. Create a third numeric variable named var_c and assign the value of var_a multiplied by var_b to it.

<!-- grade_id: cell-de9deef201587fcf, points: 0, grade: false, solution: true -->
```python

### BEGIN SOLUTION 
var_a = 2
var_b = 3 
var_c = var_a * var_b
### END SOLUTION 
```

<!-- grade_id: cell-75334516bb55a9dc, points: 1, grade: true, solution: false -->
```python
assert var_c == var_a * var_b, "var_c doesn't equal var_a * var_b"
```

## Question Two

Let's work with a vector in Python. To do so we will need to import the `numpy` package by running the following cell block. 


```python
# import numpy (dont forget to run this cell block)
import numpy as np 
```

### Part A (4 Points)

Comput the mean and sum of the following vector. Store the mean in a variable named `my_mean` and the sum in a variable named `my_sum`. 

```python
# work with the following vector (dont forget to run this line)
vec = [1,2,3,4,5]

# print values in the vector 
vec
```

Output:
```
[1, 2, 3, 4, 5]
```

<!-- grade_id: cell-653239fe29b5d894, points: 0, grade: false, solution: true -->
```python
# calculate the mean of vec using the numpy package
### BEGIN SOLUTION 
my_mean = np.mean(vec)
### END SOLUTION 

# print mean 
my_mean
```

Output:
```
3.0
```

<!-- grade_id: cell-25799336b38a26b6, points: 1, grade: true, solution: false -->
```python
# test Q2 part A 
assert my_mean == 3, "the variable my_mean doesn't equal the mean of the vector named vec"
```

### Part B 
For a moment, let's leave the vector `vec` aside. In this part, finish writing a function called `get_square()` that takes a value and returns that same value squarred. Test it by using the function to get $3^2$. Store the result in a variable called `my_nine`.

<!-- grade_id: cell-70e721899b6dd570, points: 0, grade: false, solution: true -->
```python
# define function 
def get_square(x): 
    ### BEGIN SOLUTION 
    y = x**2
    return(y)
    ### END SOLUTION 
    
# test by finding 3^2 using get_square
my_nine = get_square(3)

# print value of my_nine 
my_nine
```

Output:
```
9
```

<!-- grade_id: cell-bd53aaf6898c5269, points: 1, grade: true, solution: false -->
```python
# test Q2 part B
assert my_nine == 9, "your function must not be working because get_square(3) should return nine"
```

### Part C 
Let's return to `vec`. In this part, you will square each element in the vector `vec`. Use your function from part B to do so. This requires "intializing" a new vector where the values of the squarred elements can be stored. It then requires using a loop to iterate through each element of `vec`, square it, then store that value in the new vector. Let's called the new vector `vec_squares`. 

```python
 # initialize vec_squares (don't forget to run this)
vec_squares = [None]*5

# print intialized vector 
vec_squares
```

Output:
```
[None, None, None, None, None]
```

Remember that Pyton's indexing starts at 0, not 1. So, the first element of the vector is in the "zeroth" position. Because our vector `vec` is 5 elements long, that means we will reference the zeroth through fourth positions. When writing a for loop, the range we're iterating through (in this case) is 0 through 5 because the fifth position is the cutoff. If you're used to math notation, you can thing of `range(0,5)` as equal to $[0,5)$ where the square bracket symbolizes 0 is included, and the closed parenthese indicates all numbers **until** 5 are included. 

Fill in the loop below to get the squarred elements of `vec` and storing them in the correct positions of `vec_squares`.

<!-- grade_id: cell-3f8d1053e65e571f, points: 0, grade: false, solution: true -->
```python
# Finish the loop to square elements of vec and store them in vec_squares 
for i in range(0, 5):
    ### BEGIN SOLUTION 
    vec_squares[i] = get_square(vec[i])
    ### END SOLUTION 
    
# print vec_squares
vec_squares

np.mean(vec_squares)
```

Output:
```
11.0
```

<!-- grade_id: cell-4b36b158d7e14db1, points: 2, grade: true, solution: false -->
```python
# test Q2 Part 3 
assert np.mean(vec_squares) == 11, "Your vec_squares doesn't seem to be holding the 5 squared elements of vector named vec from part A."
```

## Question 3 (3 points)
Practice writing code from scrath! In this problem you are asked to use a for loop and and if-else statement. 

Recall our vector `vec = [1,2,3,4,5]`. Using a for loop and an if-else statement, iterate through each element of `vec`. If the value stored in the ith position is even, replace the element's value with `999`. Otherwise, leave it the same. 

<!-- grade_id: cell-20276d34b0bb4fad, points: 0, grade: false, solution: true -->
```python
# original vector 
vec = [1,2,3,4,5]

# write your for loop here to edit vec 
### BEGIN SOLUTION 
for i in range(0,5): 
    if vec[i]%2 == 0: 
        vec[i] = 999
    else: 
        vec[i] = vec[i]
### END SOLUTION 

# print an edited vec 
vec
```

Output:
```
[1, 999, 3, 999, 5]
```

<!-- grade_id: cell-fe80fdb5dc868e98, points: 3, grade: true, solution: false -->
```python
assert vec == [1, 999, 3, 999, 5], "You haven't relaced the values 2 and 4 with 999 correctly. These are in the 1st and 3rd positions (since indexing starts at 0)"
```